# Extra exercises

Resources:
- [geo-python.github.io](geo.python.github.io)
- [autogis.github.io](autogis.github.io)
- Pandas and Geopandas documentation
- Various discussion forums, such as https://stackoverflow.com/

## Import modules

Import all required modules here at the top of the notebook:

<!-- grade_id: cell-508f372d35c28ba3, points: 0, grade: false, solution: true -->
```python
# REPLACE THE ERROR BELOW WITH YOUR OWN CODE
raise NotImplementedError()
```

### List files


<div class="alert alert-info">
    
Create a function called `list_files()` that lists files based on their file type in a spesific folder. 

The function should have two input paramenters; 
-  folder path (str)
- file type extension (str)

</div>


See example function usage in the tests below.

- Define your function here: 

<!-- grade_id: cell-4f1a1455e4214365, points: 0, grade: false, solution: true -->
```python
# REPLACE THE ERROR BELOW WITH YOUR OWN CODE
raise NotImplementedError()
```

- Demonstrate the usage of the function:

```python
shapefiles = list_files("data", "shp")
print(shapefiles)
```

- Test your function by running these code cells (your function works correctly if the code cells run without errors):

<!-- grade_id: cell-c59cce9e810df0fa, points: 0, grade: true, solution: false -->
```python
# NON-EDITABLE TEST CELL

# List shapefiles in data folder
shapefiles = list_files("data", "shp")

# Check that the output is correct
assert len(shapefiles) == 1, "The function should find one shapefile from the data folder."
assert os.path.basename(shapefiles[0]) == 'MetropAccess_YKR_grid_EurefFIN.shp', "Incorrect file name."
```

<!-- grade_id: cell-599af440303e4caa, points: 0, grade: true, solution: false -->
```python
# NON-EDITABLE TEST CELL

# List text files in data folder
textfiles = list_files("data", "txt")

# Check that the output is correct
assert len(textfiles) == 7, "The function should find seven text files from the data folder."
assert any(['TravelTimes_to_5878070_Jumbo.txt' == os.path.basename(file) for file in textfiles]) , "Incorrect file list."
```

<!-- grade_id: cell-7a7714366b2ad67d, points: 0, grade: true, solution: false -->
```python
# List text files in data folder
notebooks = list_files("", "ipynb")

# Check that the output is correct
assert len(notebooks) == 1, "The function should find one Jupyter Notebook from the root folder."
assert os.path.basename(notebooks[0]) == 'Extra-tasks.ipynb', "Incorrect file name."
```

### Counting values



<div class="alert alert-info">
    
Create a function called `count_values()` that counts how many times spesific value occurs in a spesific column of a Pandas DataFrame

Then function should have three input parameters:
- input table (Pandas DataFrame)
- selected column (str)
- selected value (int)

</div>


See example function usage in the tests below.
    

- Define your function here: 

<!-- grade_id: cell-14a7f3e3a254a815, points: 0, grade: false, solution: true -->
```python
# REPLACE THE ERROR BELOW WITH YOUR OWN CODE
raise NotImplementedError()
```

- Demonstrate the usage of the function:

```python
# Read in one of the text files listed in the previous step
data = pd.read_csv(textfiles[0], sep=";")

# Count how many times -1 occurs in column "car_r_t"
print(count_values(data, "car_r_t", -1))
print(count_values(data, "pt_r_t", -1))
```

- Test your function by running these code cells (your function works correctly if the code cells run without errors):

<!-- grade_id: cell-9adac3c70ce0c3ad, points: 0, grade: true, solution: false -->
```python
# NON-EDITABLE TEST CELL

# Read in one of the text files listed in the previous step
data = pd.read_csv(textfiles[0], sep=";")

# Check correct value counts
assert count_values(data, "car_r_t", -1) == 91, "-1 occurs 91 times in column car_r_t"
assert count_values(data, "pt_r_t", -1) == 211, "-1 occurs 211 times in column pt_r_t"

```

<!-- grade_id: cell-e67d721e9da4bd33, points: 0, grade: true, solution: false -->
```python
# NON-EDITABLE TEST CELL
assert count_values(data, "car_r_t", 0) == 1, "0 occurs 1 times in column car_r_t"
assert count_values(data, "pt_r_t", 0) == 1, "0 occurs 1 times in column pt_r_t"
```

### Counting no data values

<div class="alert alert-info">
    
Create a function called `count_no_data_values()` that counts how many NO DATA values (`None` or `numpy.nan`) there are in a spesific column of a spesific table.

Then function should have two input parameters:
- input table (Pandas DataFrame)
- selected column (str)

**The function should also work IF there are no NO DATA values in the spesific column**. In this case, it should return zero (0).

</div>

See example function usage in the tests below.

- Define your function here: 

<!-- grade_id: cell-77de70ecf04a07c8, points: 0, grade: false, solution: true -->
```python
# REPLACE THE ERROR BELOW WITH YOUR OWN CODE
raise NotImplementedError()
```

- Demonstrate the usage of the function:

```python
# Read in one of the text files listed in the previous step
data = pd.read_csv(textfiles[0], sep=";")

# Replace -1 (which is the no data value in the travel time matrix) with None (or numpy.nan)
data.replace({-1 : None}, inplace=True)

# Count how many times -1 occurs in column "car_r_t"
print(count_no_data_values(data, "car_r_t"))
```

- Test your function by running these code cells (your function works correctly if the code cells run without errors):

<!-- grade_id: cell-6aa50fe75955ff91, points: 0, grade: true, solution: false -->
```python
# NON-EDITABLE TEST CELL

# Read in one of the text files listed in the previous step
data = pd.read_csv(textfiles[0], sep=";")

# Replace -1 (which is the no data value in the travel time matrix) with None (or numpy.nan)
data.replace({-1 : None}, inplace=True)

# Count how many selected values occur in selected columns
assert count_no_data_values(data, "car_r_t") == 91, "NO DATA occurs 91 times in column car_r_t"
assert count_no_data_values(data, "pt_r_t") == 211, "NO DATA occurs 211 times in column pt_r_t"
```

<!-- grade_id: cell-2bea362a26c5c372, points: 0, grade: true, solution: false -->
```python
# NON-EDITABLE TEST CELL

# Read in one of the text files listed in the previous step
data = pd.read_csv(textfiles[0], sep=";")

# NOT REPLACING -1 WITH NO DATA! 

# Count how many selected values occur in selected columns
assert count_no_data_values(data, "car_r_t") == 0, "There should be no NO DATA in this column."
assert count_no_data_values(data, "pt_r_t") == 0, "There should be no NO DATA in this column."
```

### Check projection



<div class="alert alert-info">
    
Create a function called `check_epsg()` that checks if the coordinate reference system of a geodataframe matches a specified epsg code. The function should return `True` (boolean value) if the crs matches, and `False`if it does not match.


Then function should have two input parameters:
- input layer (geopandas GeoDataFrame)
- epsg code (int)

</div>


- Define your function here: 

<!-- grade_id: cell-5feb3baf1103395b, points: 0, grade: false, solution: true -->
```python
# REPLACE THE ERROR BELOW WITH YOUR OWN CODE
raise NotImplementedError()
```

- Demonstrate the usage of the function:

```python
#Read in a shapefile
data = gpd.read_file(shapefiles[0])

# Compare data's crs with epsg 2067
check_epsg(data, 3067)
```

```python
# Compare data's crs with epsg 3857
check_epsg(data, 3857)
```

- Test your function by running these code cells (your function works correctly if the code cells run without errors):

<!-- grade_id: cell-4f8eba09960f0f15, points: 0, grade: true, solution: false -->
```python

# NON-EDITABLE TEST CELL
#Read in a shapefile
data = gpd.read_file(shapefiles[0])

# Count how many selected values occur in selected columns
assert check_epsg(data, 3067), "Data crs should match with epsg 3067"
assert not check_epsg(data, 3857), "Data crs should not match with epsg 3857"
```

### Re-project data







<div class="alert alert-info">
    
Create a function called `reproject_layer()` that does the following:
    
- checks if the input layer has a crs definition. 
    - If no crs is defined (`layer.crs=""` or `layer.crs=None`) , the function throws an assertion error and stops the process.
- checks if the crs of the input layer matches with the specified epsg code
    - if the crs does not match, the input layer is reprojected into the specified epsg and the function returns the re-projected layer
    - if the crs matches with the specified epsg code, then the function returns the original input layer (you can also add an informative message about this)


Then function should have two input parameters:
- input layer (geopandas GeoDataFrame)
- epsg code (int)

**Note: you can use `check_epsg()`function inside this function**

</div>


- Define your function here: 

<!-- grade_id: cell-b88dfde59260c965, points: 0, grade: false, solution: true -->
```python
# REPLACE THE ERROR BELOW WITH YOUR OWN CODE
raise NotImplementedError()
```

- Demonstrate the usage of the function:

```python
#Read in a shapefile
data = gpd.read_file(shapefiles[0])

# Reproject the layer to epsg:3857 using the function
reprojected_data = reproject_layer(data, 3857)
reprojected_data.head()
```

<!-- grade_id: cell-3199f13baa24e814, points: 0, grade: true, solution: false -->
```python
# NON-EDITABLE TEST CELL

#Read in a shapefile
data = gpd.read_file(shapefiles[0])

# Use the function to it's own epsg
result = reproject_layer(data, 3067)

# Check that original and result geometries match
assert all(result["geometry"] == data["geometry"]), "Values should be the same!"
```

<!-- grade_id: cell-61e276ab0476982c, points: 0, grade: true, solution: false -->
```python
# NON-EDITABLE TEST CELL

#Read in a shapefile
data = gpd.read_file(shapefiles[0])

# Set data crs as None
data.crs = None

# Try to apply the function, error message will be printed on the screen if the process fails
try: 
    result = reproject_layer(data, 3067)

except AssertionError as msg:
    print(msg)
```

<!-- grade_id: cell-232d59fe1dc02687, points: 0, grade: true, solution: false -->
```python
# NON-EDITABLE TEST CELL
#Read in a shapefile
data = gpd.read_file(shapefiles[0])

# Set data crs as None
data.crs = ""

# Try to apply the function
try: 
    
    result = reproject_layer(data, 3067)

except AssertionError as msg:
    # Error message will be printed on the screen if the process fails
    print(msg)
```