# Extra exercises 2

### Simple xmas tree

Create a function `make_xmas_tree()`that plots a spruce:


         *
        ***
       *****
      *******
     *********
     
The function should be able to create different size trees. Input parameter three should produce this:
   
       *
      ***
     *****
     

The function should have one input paramenter (int) that specifies how many rows the tree has. There should be no whitespace on the bottom row.

- Define your function here: 

<!-- grade_id: cell-66948068ef06bb07, points: 0, grade: false, solution: true -->
```python
# REPLACE THE ERROR BELOW WITH YOUR OWN CODE
raise NotImplementedError()
```

- Demonstrate the usage of the function:

```python
small_spruce = make_xmas_tree(3)
print(small_spruce)
```

```python
big_spruce = make_xmas_tree(30)
print(big_spruce)
```

- Test your function by running these code cells (your function works correctly if the code cells run without errors):

<!-- grade_id: cell-11e68e7fcfea740c, points: 0, grade: true, solution: false -->
```python
# Create a spruce
small_spruce = make_xmas_tree(3)

# Check the result
assert small_spruce ==  '  *\n ***\n*****\n', "wrong kind of xmas tree!"
```

```python
# Create a spruce
medium_spruce = make_xmas_tree(12)

# Check the result
assert medium_spruce == '           *\n          ***\n         *****\n        *******\n       *********\n      ***********\n     *************\n    ***************\n   *****************\n  *******************\n *********************\n***********************\n', "Wrong kind of xmas tree!"
```